## Habilidades en práctica

Al realizar este taller podrás revisar tu progreso para:

**1.** Utilizar arreglos y operaciones básicas en `numpy`. <br>
**2.** Crear, indexar y utilizar métodos para explorar y manipular `Series` y `DataFrame` en `pandas`.

## Instrucciones

En cada uno de los siguientes ejercicios deberás escribir el código solicitado estrictamente en las celdas indicadas para ello, teniendo en cuenta las siguientes recomendaciones:

* No crear, eliminar o modificar celdas de este Notebook (salvo lo que se te indique), pues puede verse afectado el proceso de calificación automática.
* La calificación se realiza de manera automática con datos diferentes a los proporcionados en este taller. Por consiguiente, tu código debe funcionar para diferentes instancias de cada uno de los ejercicios; una instancia hace referencia al valor de los parámetros.
* La calificación de cada ejercicio depende del valor que tome la variable especificada en su enunciado. Por lo tanto, aunque declares variables adicionales en tu código, es escencial que utilices los nombres de las variables propuestas en los enunciados de los ejercicios para guardar el resultado final.

Esta es una actividad calificada y, por lo tanto, debe ser resuelta individualmente.

A continuación puedes ver un ejemplo de los ejercicios que encontrarás en este taller.

## Ejercicios
En la siguente celda encuentras declarados los paquetes necesarios para el desarollo de este taller.

```python
# Esta celda no es modificable.

import numpy as np
import pandas as pd
```

### Ejemplo
A continuación encuentras dos arreglos. Cada uno contiene los salarios (en USD) de los funcionarios de un área de una empresa.

```python
# Esta celda no es modificable.

mercadeo = np.array([2000, 1500, 1300, 1600])
logistica = np.array([1800, 1600, 1200, 1100])
```

Implementa una función llamada `suma` que reciba dos arreglos por parámetro, los una y retorne la suma de los elementos de la unión.

*Esta función debe retornar un `int`.*

```python
# Implementa tu respuesta en esta celda.

def suma(arreglo1, arreglo2):

    concatenado = np.concatenate((arreglo1,arreglo2))
    respuesta = concatenado.sum()

    return respuesta

suma(mercadeo, logistica)
```

### Ejercicio 1
A continuación encuentras dos arreglos. Cada uno contiene los salarios (en USD) de los funcionarios de un área de una empresa.

```python
# Esta celda no es modificable.
mercadeo = np.array([2000, 1500, 1300, 1600])
logistica = np.array([1800, 1600, 1200, 1600])

```

Implementa una función llamada `promedio` que reciba dos arreglos por parámetro, los una y retorne el promedio de los elementos de la unión.

*Esta función debe retornar un `float`*.

<!-- grade_id: ej1-sol, points: 0, grade: false, solution: true -->
```python
# your code here
# toca instalar numpy para que funcione en Colab por eso lo hago profe    
!pip install numpy
import numpy as np
def promedio(arreglo1, arreglo2):
    concatenado = np.concatenate((arreglo1, arreglo2))
    return concatenado.mean()
mercadeo = np.array([2000, 1500, 1300, 1600])
logistica = np.array([1800, 1600, 1200, 1600])
print(promedio(mercadeo, logistica)) 



```

Output:
```
Requirement already satisfied: numpy in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (2.2.6)
1575.0

[notice] A new release of pip is available: 23.0.1 -> 25.2
[notice] To update, run: python.exe -m pip install --upgrade pip
```

<!-- grade_id: ej1-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
mercadeo_test = np.array([2000, 1800, 1300])
logistica_test = np.array([1100, 1900, 1600, 1500])
# Respuestas 1800 og  y 1600 test

# Caso 1: no existe la función.
try:
    promedio
    assert type(promedio) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada promedio.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    promedio(mercadeo, logistica)
    promedio(mercadeo_test, logistica_test)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un número
assert type(promedio(mercadeo_test, logistica_test)) == np.float64, f"Tu función debe retornar un valor de tipo {np.float64}."

# Caso 4: respuesta explicita
assert promedio(mercadeo_test, logistica_test) != 1800.0, "Tu respuesta es incorrecta para una instancia diferente. Utiliza los parámetros."

# Caso 5: procedimiento incorrecto
assert promedio(mercadeo_test, logistica_test) <= 1600.0, "Tu respuesta es incorrecta. Asegurate de cocatenar los arreglos y luego calcular un promedio."

# Caso 6: retorna un numero distinto del esperado
assert promedio(mercadeo_test, logistica_test) == 1600.0, "Tu respuesta es incorrecta. Es posible que hayas omitido algún valor al realizar el cálculo."
assert promedio(mercadeo, logistica) == 1575.0, "Tu respuesta es incorrecta. Es posible que hayas omitido algún valor al realizar el cálculo."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 2

A continuación encuentras un `Series` que contiene la edad de Sonia y sus hermanos. Ellos quieren quienes son los dos mayores para estar atentos a la vacunación.

```python
# Esta celda no es modificable.

hermanos = pd.Series({"Sonia": 40, "Patricia": 42, "Jorge": 37, "Alejandro": 45, "Aura": 47, "Carlos": 40, "Claudia": 41, "Victor": 43})
```

Implementa una función llamada `mayores` que reciba un `Series` por parámetro y retorne los dos mayores valores en una tupla.

*Esta función debe retornar una tupla de la forma* `(mayor_valor1, mayor_valor2)` *, dónde* `mayor_valor1` *es mayor o igual que* `mayor_valor2` *.*

<!-- grade_id: ej2-sol, points: 0, grade: false, solution: true -->
```python
# your code here

!pip install pandas

import pandas as pd


hermanos = pd.Series({
    "Sonia": 40, "Patricia": 42, "Jorge": 37,
    "Alejandro": 45, "Aura": 47, "Carlos": 40,
    "Claudia": 41, "Victor": 43
})


def mayores(series):
    mayor1 = -1
    mayor2 = -1
    for edad in series:
        if edad > mayor1:
            mayor2 = mayor1
            mayor1 = edad
        elif edad > mayor2:
            mayor2 = edad
    return (mayor1, mayor2)




```

Output:
```
Requirement already satisfied: pandas in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (2.3.2)
Requirement already satisfied: tzdata>=2022.7 in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (from pandas) (2025.2)
Requirement already satisfied: pytz>=2020.1 in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (from pandas) (2025.2)
Requirement already satisfied: python-dateutil>=2.8.2 in c:\users\jdavi\appdata\roaming\python\python310\site-packages (from pandas) (2.9.0.post0)
Requirement already satisfied: numpy>=1.22.4 in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (from pandas) (2.2.6)
Requirement already satisfied: six>=1.5 in c:\users\jdavi\appdata\roaming\python\python310\site-packages (from python-dateutil>=2.8.2->pandas) (1.17.0)

[notice] A new release of pip is available: 23.0.1 -> 25.2
[notice] To update, run: python.exe -m pip install --upgrade pip
```

<!-- grade_id: ej2-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
hermanos_test = pd.Series({"Sonia": 43, "Patricia": 46, "Jorge": 39, "Alejandro": 40, "Aura": 30, "Carlos": 41, "Claudia": 40})

# Caso 1: no existe la función.
try:
    mayores
    assert type(mayores) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada mayores.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    mayores(hermanos)
    mayores(hermanos_test)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna una lista
assert type(mayores(hermanos_test)) == type(("",)), f"Tu función debe retornar un objeto de tipo {type(tuple).__name__}."

# Caso 4: tamaño incorrecto
assert len(mayores(hermanos_test)) == 2, "Tu función retorna una lista de un tamaño diferente al solicitado."

# Caso 5: procedimiento incorrecto
assert mayores(hermanos_test) != [47, 45], "Tu respuesta es incorrecta para una instancia diferente. Utiliza los parámetros."

# Caso 6: la lista no contiene todos los valores correctos
try:
    assert mayores(hermanos_test).count(46) >= 1
    assert mayores(hermanos_test).count(43) >= 1
except AssertionError as e:
    e.args += ("Tu respuesta no contiene la información solicitada.",)
    raise e

# Caso 7: la lista es estrictamente correcta
try:
    assert mayores(hermanos_test)[0] == 46
    assert mayores(hermanos_test)[1] == 43
except AssertionError as e:
    e.args += ("Tu respuesta tiene un orden incorrecto.",)
    raise e

# Caso 8: retorna un numero distinto del esperado
assert mayores(hermanos) == (47, 45), "Tu respuesta es incorrecta."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

## Salarios en latinoamérica

A continuación, encuentras declarado un `DataFrame` que contiene la edad, la ciudad de residencia y el salario (en USD) de nombres de algunos empleados de una multinacional. Utiliza esta información para dar respuesta a los ejercicios 3, 4 y 5.

```python
# Esta celda no es modificable.

e1 = [ 28,"Bogotá", 2000]
e2 = [ 37, "Lima", 3200 ]
e3 = [ 23,"Bogotá", 1700]
e4 = [ 25, "Buenos Aires", 1100 ]
e5 = [ 43,"Buenos Aires", 3300]
e6 = [ 58, "Lima", 5500 ]
e7 = [ 32,"Bogotá", 2700]
e8 = [ 35, "Quito", 2500 ]

empleados = pd.DataFrame([e1,e2, e3, e4, e5, e6, e7, e8],
                         index = ["Carlos", "David", "Ana", "Maria", "Felipe", "Luisa", "Juan", "Camila"],
                         columns = ["Edad", "Ciudad Residencia", "Salario (en USD)"])
empleados
```

### Ejercicio 3

Implementa una función llamada `ciudades` que reciba un `DataFrame` por parámetro y retorne el número de ciudades diferentes en las que residen los empleados de la multinacional.

*Esta función debe retornar un `int`*.

<!-- grade_id: ej3-sol, points: 0, grade: false, solution: true -->
```python
# your code here

!pip install pandas


import pandas as pd
e1 = [28,"Bogotá", 2000]
e2 = [37,"Lima", 3200]
e3 = [23,"Bogotá", 1700]
e4 = [25,"Buenos Aires", 1100]
e5 = [43,"Buenos Aires", 3300]
e6 = [58,"Lima", 5500]
e7 = [32,"Bogotá", 2700]
e8 = [35,"Quito", 2500]


empleados = pd.DataFrame(
    [e1,e2,e3,e4,e5,e6,e7,e8],
    index=["Carlos", "David", "Ana", "Maria", "Felipe", "Luisa", "Juan", "Camila"],
    columns=["Edad", "Ciudad Residencia", "Salario (en USD)"]
)


def ciudades(df):
    lista_ciudades = df["Ciudad Residencia"]
    distintas = set(lista_ciudades)
    return len(distintas)


 


```

Output:
```
Requirement already satisfied: pandas in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (2.3.2)
Requirement already satisfied: python-dateutil>=2.8.2 in c:\users\jdavi\appdata\roaming\python\python310\site-packages (from pandas) (2.9.0.post0)
Requirement already satisfied: tzdata>=2022.7 in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (from pandas) (2025.2)
Requirement already satisfied: numpy>=1.22.4 in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (from pandas) (2.2.6)
Requirement already satisfied: pytz>=2020.1 in c:\users\jdavi\appdata\local\programs\python\python310\lib\site-packages (from pandas) (2025.2)
Requirement already satisfied: six>=1.5 in c:\users\jdavi\appdata\roaming\python\python310\site-packages (from python-dateutil>=2.8.2->pandas) (1.17.0)

[notice] A new release of pip is available: 23.0.1 -> 25.2
[notice] To update, run: python.exe -m pip install --upgrade pip
```

<!-- grade_id: ej3-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Test paremters
e1 = [ 28,"Quito", 2000]
e2 = [ 37, "Lima", 3200 ]
e3 = [ 22,"Bogotá", 1700]
e4 = [ 35, "Santiago", 1100]
e5 = [ 41,"Buenos Aires", 3900]
e6 = [ 58, "Lima", 5500]
e7 = [ 22,"Brasilia", 2100]
e8 = [ 35, "Bogotá", 2500]

empleados_test = pd.DataFrame([e1,e2, e3, e4, e5, e6, e7, e8],
                              index = ["Carlos", "David", "Ana", "Maria", "Felipe", "Luisa", "Juan", "Camila"],
                              columns = ["Edad", "Ciudad Residencia", "Salario (en USD)"])

# Caso 1: no existe la función.
try:
    ciudades
    assert type(ciudades) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada ciudades.")

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    ciudades(empleados)
    ciudades(empleados_test)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un int
assert type(ciudades(empleados_test)) == int, f"Tu función debe retornar un objeto de tipo {int.__name__}."

# Caso 4: respuesta explicita
assert ciudades(empleados_test) != 4, "Tu respuesta es incorrecta para una instancia diferente. Utiliza los parámetros."

# Caso 5: Retorna el número de elementos de la columna 'Ciudad Residencia'
assert ciudades(empleados_test) != 8, "Tu respuesta es el número de elementos de la columna 'Ciudad Residencia'."

# Caso 6: retorna un numero distinto del esperado
assert ciudades(empleados_test) == 6, "Tu respuesta es incorrecta."
assert ciudades(empleados) == 4, "Tu respuesta es incorrecta."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 4

Implementa una función llamada `mayor_al_promedio` que reciba un `DataFrame` y un índice de este `DataFrame` por parámetro, y retorne `True` si el salario asociado a este índice es mayor al promedio o `False` de lo contrario.

*Esta función debe retornar* `True` *o* `False` *.*

<!-- grade_id: ej4-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def mayor_al_promedio(df, indice):
    promedio = df["Salario (en USD)"].mean()
    salario = df.loc[indice, "Salario (en USD)"]
    return bool(salario > promedio)

```

<!-- grade_id: ej4-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Test paremters
e1 = [ 28,"Quito", 2000]
e2 = [ 37, "Lima", 3200 ]
e3 = [ 22,"Bogotá", 1700]
e4 = [ 35, "Santiago", 1100 ]
e5 = [ 41,"Buenos Aires", 3900]
e6 = [ 58, "Lima", 5500 ]
e7 = [ 22,"Brasilia", 2100]
e8 = [ 35, "Bogotá", 2500 ]

empleados_test = pd.DataFrame([e1,e2, e3, e4, e5, e6, e7, e8],
                              index=["Carlos", "David", "Ana", "Maria", "Felipe", "Luisa", "Juan", "Camila"],
                              columns=["Edad", "Ciudad Residencia", "Salario (en USD)"])

# Caso 1: no existe la función.
try:
    mayor_al_promedio
    assert type(mayor_al_promedio) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada mayor_al_promedio.")

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    mayor_al_promedio(empleados, 'Juan')
    mayor_al_promedio(empleados_test, 'Luisa')
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un boolean
assert type(mayor_al_promedio(empleados_test, 'Luisa')) == type(True), f"Tu función debe retornar un objeto de tipo {type(True).__name__}."

# Caso 4: procedimiento incorrecto
assert mayor_al_promedio(empleados_test, 'Luisa') != False, "Tu respuesta es incorrecta para una instancia diferente. Utiliza los parámetros."

# Caso 5: retorna un numero distinto del esperado
assert mayor_al_promedio(empleados_test, 'Luisa') == True, "Tu respuesta es incorrecta."
assert mayor_al_promedio(empleados, 'Juan') == False, "Tu respuesta es incorrecta."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 5

Implementa una función llamada `modificar` que reciba un `DataFrame` por parámetro, elimine la segunda fila, cambie el nombre de la columna 'Salario (en USD)' a 'Sueldo (en USD)' y retorne un nuevo `DataFrame` con las modificaciones realizadas.

*Esta función debe retornar un* `DataFrame` *.*

<!-- grade_id: ej5-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def modificar(df):
    nuevo_df = df.drop(df.index[1])
    nuevo_df = nuevo_df.rename(columns={"Salario (en USD)": "Sueldo (en USD)"})
    return nuevo_df

```

<!-- grade_id: ej5-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Test paremters
e1 = [ 28,"Quito", 2000]
e2 = [ 37, "Lima", 3200 ]
e3 = [ 22,"Bogotá", 1700]
e4 = [ 35, "Santiago", 1100 ]
e5 = [ 41,"Buenos Aires", 3900]
e6 = [ 58, "Lima", 5500 ]
e7 = [ 22,"Brasilia", 2100]
e8 = [ 35, "Bogotá", 2500 ]

empleados_test = pd.DataFrame([e1,e2, e3, e4, e5, e6, e7, e8],
                              index=["Carlos", "Diego", "Ariel", "Maria", "Felipe", "Luisa", "Juan", "Camila"],
                              columns=["Edad", "Ciudad Residencia", "Salario (en USD)"])

# Caso 1: no existe la función.
try:
    modificar
    assert type(modificar) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada modificar.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    modificar(empleados)
    modificar(empleados_test)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un dataframe
assert type(modificar(empleados_test)) == type(empleados_test), f"Tu función debe retornar un objeto de tipo {type(empleados_test).__name__}."

# Caso 4: Cambio de nombre correcto
assert modificar(empleados_test).columns[2] == 'Sueldo (en USD)', "Tu respuesta es incorrecta. Recuerda modificar el nombre de la columna 'Salario (en USD)'"

# Caso 5: Elimino la fila
modificar(empleados_test).index[[1]] == 'Ariel', "Tu respuesta es incorrecta. Recuerda eliminar la segunda fila del DataFrame."

# Caso 6: los cambios son estrictamente correctos
try:
    assert modificar(empleados).columns[2] == 'Sueldo (en USD)'
    assert modificar(empleados).index[[1]] == 'Ana'
except AssertionError as e:
    e.args += ("Tu respuesta es incorrecta.",)
    raise e

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

## Lista de países

A continuación encuentras un `DataFrame` con un listado de países con su respectivo nombre en inglés y en español.

```python
# Esta celda no es modificable.

paises = pd.read_csv("./Archivos/DiccionarioPaises.csv", encoding = "UTF-8", delimiter = ";")
paises.head()
```

### Ejercicio 6

Implementa una función llamada `dif_longitud` que reciba un `DataFrame` por parámetro, elimine los países cuya traducción cambie en longitud  y retorne un `DataFrame` con las modificaciones realizadas.

*Esta función debe retornar un* `DataFrame` *.*

<!-- grade_id: ej6-sol, points: 0, grade: false, solution: true -->
```python
# your code here
import pandas as pd
import os
print("Carpeta actual:", os.getcwd())
paises = pd.read_csv(r"C:/Users/jdavi/Tutorial/Archivos/DiccionarioPaises.csv", encoding="UTF-8", delimiter=";")
def dif_longitud(df):
    return df[df["País en inglés"].str.len() == df["País en español"].str.len()]
print(paises.head())
print(dif_longitud(paises).head())



```

Output:
```
Carpeta actual: d:\documentos\GitHub\ELECTIVA-PROFESIONAL-III\Semana3\Taller libreria para menejo de datos
File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\parsers\readers.py:1880, in TextFileReader._make_engine(self, f, engine)
   1878     if "b" not in mode:
   1879         mode += "b"
-> 1880 self.handles = get_handle(
   1881     f,
   1882     mode,
   1883     encoding=self.options.get("encoding", None),
   1884     compression=self.options.get("compression", None),
   1885     memory_map=self.options.get("memory_map", False),
   1886     is_text=is_text,
   1887     errors=self.options.get("encoding_errors", "strict"),
   1888     storage_options=self.options.get("storage_options", None),
   1889 )
   1890 assert self.handles is not None
   1891 f = self.handles.handle

File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\common.py:873, in get_handle(path_or_buf, mode, encoding, compression, memory_map, is
```

<!-- grade_id: ej6-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
paises = pd.read_csv("./Archivos/DiccionarioPaises.csv", encoding = "UTF-8", delimiter = ";")
nombre_columnas = ['País en español', 'País en inglés']

# Caso 1: no existe la función.
try:
    dif_longitud
    assert type(dif_longitud) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada dif_longitud.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    dif_longitud(paises)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame.
assert type(dif_longitud(paises)) == type(paises), f"Tu función debe retornar un objeto de tipo {type(paises).__name__}."

# Caso 4: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*dif_longitud(paises).columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no coinciden con los de los datos del archivo. Asegurate de estar filtrando las columnas indicadas."

# Caso 5: devuelve un dataframe con cantidad de filas errada.
assert dif_longitud(paises).shape[0] == 127, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 6: No elimina las filas correctas
assert sum((dif_longitud(paises)['País en español']+dif_longitud(paises)['País en inglés']).apply(len) % 2) == 0, "Tu función no elimina las filas correctas del DataFrame."


# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\parsers\readers.py:1880, in TextFileReader._make_engine(self, f, engine)
   1878     if "b" not in mode:
   1879         mode += "b"
-> 1880 self.handles = get_handle(
   1881     f,
   1882     mode,
   1883     encoding=self.options.get("encoding", None),
   1884     compression=self.options.get("compression", None),
   1885     memory_map=self.options.get("memory_map", False),
   1886     is_text=is_text,
   1887     errors=self.options.get("encoding_errors", "strict"),
   1888     storage_options=self.options.get("storage_options", None),
   1889 )
   1890 assert self.handles is not None
   1891 f = self.handles.handle

File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\common.py:873, in get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text, errors, storage_options)
    868 elif isinstance(handle, str):
    869     # Check whether the filen
```

### Ejercicio 7

Implementa una función llamada `paises_por_a` que reciba un `DataFrame` por parámetro y retorne el mismo `DataFrame` con los países cuyo nombre en español inicia con la letra `"A"`.

*Esta función debe retornar un `DataFrame`*.

<!-- grade_id: ej7-sol, points: 0, grade: false, solution: true -->
```python
# your code here

def paises_por_a(df):
    resultado = df[df["País en español"].str.startswith("A")]
    return resultado

```

<!-- grade_id: ej7-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
paises = pd.read_csv("./Archivos/DiccionarioPaises.csv", encoding = "UTF-8", delimiter = ";")
nombre_columnas = ['País en español', 'País en inglés']

# Caso 1: no existe la función.
try:
    paises_por_a
    assert type(paises_por_a) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada paises_por_a.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    paises_por_a(paises)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un dataframe.
assert type(paises_por_a(paises)) == type(paises), f"Tu función debe retornar un objeto de tipo {type(paises).__name__}."

# Caso 4: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado
assert [*paises_por_a(paises).columns] == nombre_columnas, "Tu función retorna un DataFrame con columnas cuyos nombres no coinciden con los de los datos del archivo. Asegurate de estar filtrando las columnas indicadas."

# Caso 5: devuelve un dataframe con cantidad de filas errada.
assert paises_por_a(paises).shape[0] == 13, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 6: No elimina las filas correctas
lista_a = ['Afganistán', 'Albania', 'Alemania', 'Andorra', 'Angola', 'Antigua y Barbuda', 'Arabia Saudí', 'Argelia', 'Argentina', 'Armenia', 'Australia', 'Austria', 'Azerbaiyán']

assert lista_a == [*paises_por_a(paises)['País en español']], "Tu función no elimina las filas correctas del DataFrame."


# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\parsers\readers.py:1880, in TextFileReader._make_engine(self, f, engine)
   1878     if "b" not in mode:
   1879         mode += "b"
-> 1880 self.handles = get_handle(
   1881     f,
   1882     mode,
   1883     encoding=self.options.get("encoding", None),
   1884     compression=self.options.get("compression", None),
   1885     memory_map=self.options.get("memory_map", False),
   1886     is_text=is_text,
   1887     errors=self.options.get("encoding_errors", "strict"),
   1888     storage_options=self.options.get("storage_options", None),
   1889 )
   1890 assert self.handles is not None
   1891 f = self.handles.handle

File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\common.py:873, in get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text, errors, storage_options)
    868 elif isinstance(handle, str):
    869     # Check whether the filen
```

### Ejercicio 8

A continuación encuentras una tupla con las letras del alfabeto.

```python
# Esta celda no es modificable.

alfabeto = ('A', 'B', 'C', 'D', 'E', 'F', 'G', 'H', 'I',
            'J', 'K', 'L', 'M', 'N', 'Ñ', 'O', 'P', 'Q',
            'R', 'S', 'T', 'U', 'V', 'W', 'X', 'Y', 'Z')
```

Implementa una función llamada `inicia_con` que reciba un `DataFrame` por parámetro y retorne un diccionario cuyas llaves sean las letras del abecedario y el valor de cada llave sea el número de países cuyo nombre en inglés inicia con la respectiva letra.

*Esta función debe retornar un diccionario*.

<!-- grade_id: ej8-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def inicia_con(df):
    resultado = {}
    
    for letra in alfabeto:
        contador = 0
        for pais in df["País en inglés"]:
            if pais.startswith(letra):
                contador += 1
        
        resultado[letra] = contador
    
    return resultado

```

<!-- grade_id: ej8-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
paises = pd.read_csv("./Archivos/DiccionarioPaises.csv", encoding = "UTF-8", delimiter = ";")
nombre_columnas = ['País en español', 'País en inglés']
alfabeto = ('A', 'B', 'C', 'D', 'E', 'F', 'G', 'H', 'I',
            'J', 'K', 'L', 'M', 'N', 'Ñ', 'O', 'P', 'Q',
            'R', 'S', 'T', 'U', 'V', 'W', 'X', 'Y', 'Z')

# Caso 1: no existe la función.
try:
    inicia_con
    assert type(inicia_con) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada inicia_con.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    inicia_con(paises)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un diccionario.
assert type(inicia_con(paises)) == type({}), f"Tu función debe retornar un objeto de tipo '{type({}).__name__}'."

# Caso 4: Llaves son el abecedario
assert tuple(inicia_con(paises).keys()) == alfabeto, "Tu diccionario tiene llaves diferentes al abecedario."

# Caso 5: Comparamos las que son iguales a cero y otros valores.
try:
    assert (inicia_con(paises).get('Ñ')) == 0
    assert (inicia_con(paises).get('X')) == 0
    assert (inicia_con(paises).get('B')) == 18
    assert (inicia_con(paises).get('W')) == 1
    assert (inicia_con(paises).get('S')) == 24

except AssertionError as e:
    e.args += ("Tu respuesta es incorrecta. Revisa tu procedimiento",)
    raise e

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\parsers\readers.py:1880, in TextFileReader._make_engine(self, f, engine)
   1878     if "b" not in mode:
   1879         mode += "b"
-> 1880 self.handles = get_handle(
   1881     f,
   1882     mode,
   1883     encoding=self.options.get("encoding", None),
   1884     compression=self.options.get("compression", None),
   1885     memory_map=self.options.get("memory_map", False),
   1886     is_text=is_text,
   1887     errors=self.options.get("encoding_errors", "strict"),
   1888     storage_options=self.options.get("storage_options", None),
   1889 )
   1890 assert self.handles is not None
   1891 f = self.handles.handle

File c:\Users\jdavi\AppData\Local\Programs\Python\Python310\lib\site-packages\pandas\io\common.py:873, in get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text, errors, storage_options)
    868 elif isinstance(handle, str):
    869     # Check whether the filen
```

## Créditos

**Autor(es):** Camilo Falla Moreno, Juan David Reyes Jaimes, Alejandro Mantilla Redondo, Diego Alejandro Cely Gómez

**Fecha última actualización:** 17/09/2021